# COMP90016 - Practice Assignment



Version 1. Last edited 13/03/2023.

```python
### Fill in your student details here
NAME = "Rafsan Al Mamun"
ID = "1407776"
```

## Background
This practice assignment has been developed so you can trial the assignment submision process. Please submit your completed assignment through turnitin on the LMS as both a Jupyter notebook file and an HTML file. Note that no other files or formats will be accepted.

## Marking

**This practice assignment is not worth any marks. It does not contribute to your final grade.**

Cells that must be completed to receive marks are clearly labelled. Some cells are code cells, in which you must complete the code to solve a problem. Others are markdown cells, in which you must write your answers to short-answer questions. 

Cells that must be completed to receive marks are labelled like this:

`# -- GRADED CELL (1 mark) - complete this cell --`

### Completing code cells

- You will see the following text in graded code cells:

``` python
# YOUR CODE HERE
raise NotImplementedError()
```

- ***You must remove the `raise NotImplementedError()` line from the cell, and replace it with your solution.***


- Only add answers to graded cells. If you want to import a library or use a helper function, this must be included in a graded cell.


- Include code comments in your solutions. Well commented code can help you to receive partial marks even if the final solution is incorrect. 


## Submitting this assignment

Before you turn this assignment in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you have filled in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE"


Your completed notebook file containing all your answers must be turned in via LMS in `.ipynb` format.
You must also submit a copy of this notebook in `html` format with the output cleared.
You can do this by using the `clear all output` option in the menu.

Your submission should include **only two** files with names formatted as: **Practice_1.ipynb** and **Practice_1.html**


## Task 1 - Squares

<div style="color: rgb(27,94,32); background: rgb(200,230,201); border: solid 1px rgb(129,199,132); padding: 10px;">
<h3>Question 1</h3>

<b>Challange:</b> Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.
    
- [ ] Take as input an int n
- [ ] Return a list of squares from 1 to n
- [ ] Return a value error if n < 1    
    
</div>


<!-- grade_id: Q-1, points: 0, grade: false, solution: true -->
```python
# -- GRADED CELL (2 marks) - complete this cell --
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: T-1-1, points: 1, grade: true, solution: false -->
```python
# --- AUTOGRADING CELL DO NOT EDIT ----

"""Check that squares returns the correct output for several inputs"""

assert squares(1) == [1]
assert squares(2) == [1, 4]
assert squares(10) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100]
assert squares(11) == [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]

```

<!-- grade_id: T-1-2, points: 1, grade: true, solution: false -->
```python
# --- AUTOGRADING CELL DO NOT EDIT ----

"""Check that squares raises an error for invalid inputs"""
try:
    squares(0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

try:
    squares(-4)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

```



<div style="color: rgb(0,96,100); background: rgb(178,235,242); border: solid 1px rgb(77,208,225); padding: 10px;">

<h3>Question 2:</h3>

Adam has to manually fix ridiculous file names like `Final-final-draft_Kevin's Assignment 1 (0) (1) (1)<ID =9771827>.ipynb` so that they don't break the grading software. This does not spark joy.

<b>Question:</b> Reviewing the **Submitting this Assignment** section above - what are the names of the only two files you need to submit for this practice assignment?
    
Answer in the cell below. 

**2 marks**, maximum of 50 words
 
</div>




#### -- GRADED CELL (2 marks) - complete this cell -- 

YOUR ANSWER HERE